Question 1: write an algorithm that adds three numbers

<!-- grade_id: cell-865db9526f3139e1, points: 0, grade: false, solution: true -->
```c
// A test program that we wish to run
#include <stdio.h>
#include <assert.h>

int main () {
    // for this question, the following variable names must be used:
    int x; // the first number to be added (input)
    int y; // the second number to be added (input)
    int z; // the third number to be added (input)
    int out; // the result of the calculation
    // [PRECONDITION TAG] 
    
    out = x;
    out -= y;
    out += z;
    
    // [POSTCONDITION TAG]
    return 0;
}
```

<!-- grade_id: cell-d496c5c0af6c4b0d, points: 2, grade: true, solution: false -->
```c
#include <stdio.h>
#include <stdlib.h>
#include <string.h>
#include <assert.h>

const char* cellID = "cell-865db9526f3139e1"; // Enter the cell ID number here
const char* aName = " TestAssignment.ipynb"; // Be sure to leave a space prepended to the filename
// Injected code must follow some rules:
// preconditions must be space separated, with semicolons omitted.
// Prepend one space
// omit all spaces that do not separate preconditions.  
const char* preconditions = " x=4 y=5 z=7"; // Code to be injected at precondition tag (No Whitespace!)
const char* postcondition = " out==16"; // Expression in terminal assert statement (No Whitespace!)

int main() {
    char cmd [200] = "./AutogradeCell.sh ";
    strcat(cmd, cellID);
    strcat(cmd, aName);
    strcat(cmd, postcondition);
    strcat(cmd, preconditions);
    int status = system(cmd);
    assert (status == 0);
    printf("Test Case Passed!");
    system("./cleanup.sh");
}
```

Output:
```
Testing Cell #: cell-865db9526f3139e1
In Notebook: TestAssignment.ipynb
['scoop.py', 'cell-865db9526f3139e1', 'TestAssignment.ipynb', 'out==16', 'x=4', 'y=5', 'z=7']
Target Acquired.
Injecting Postcondition... out==16
Injecting Preconditions... ['x=4', 'y=5', 'z=7']
Code compiles!
answer: answer.c:19: main: Assertion `out==16' failed.
Code Executes!
Test Case Passed!./AutogradeCell.sh: line 13:  2236 Aborted                 ./answer
Cleanup Complete!
```

```c

```